# SLU 17 - Exam Prep I

## Batch 2 - Wave 2 Python Exam
This is the Python exam from batch 2 - wave 2. Exam Duration: 2h

```python
import json
from operator import itemgetter
import random
#used for evaluation
import utils
```

```python
data=utils.generate_data()
```

```python
data
```

Output:
```
[{'animal_name': 'gull',
  'hair': 0,
  'feathers': '?',
  'eggs': 1,
  'milk': 0,
  'aquatic': '?',
  'toothed': 0,
  'fins': 0,
  'legs': 2,
  'class_type': 'Bird'},
 {'animal_name': 'hare',
  'hair': 1,
  'feathers': '0',
  'eggs': 0,
  'milk': 1,
  'aquatic': '0',
  'toothed': 1,
  'fins': 0,
  'legs': 4,
  'class_type': 'Mammal'},
 {'animal_name': 'gull',
  'hair': 0,
  'feathers': '?',
  'eggs': 1,
  'milk': 0,
  'aquatic': '?',
  'toothed': 0,
  'fins': 0,
  'legs': 2,
  'class_type': 'Bird'},
 {'animal_name': 'pheasant',
  'hair': 0,
  'feathers': '?',
  'eggs': 1,
  'milk': 0,
  'aquatic': '0',
  'toothed': 0,
  'fins': 0,
  'legs': 2,
  'class_type': 'Bird'},
 {'animal_name': 'cavy',
  'hair': 1,
  'feathers': '0',
  'eggs': 0,
  'milk': 1,
  'aquatic': '0',
  'toothed': 1,
  'fins': 0,
  'legs': 4,
  'class_type': 'Mammal'},
 {'animal_name': 'haddock',
  'hair': 0,
  'feathers': '0',
  'eggs': 1,
  'milk': 0,
  'aquatic': '?',
  'toothed': 1,
  'fins': 1,
  'legs': 0,
  'cla
```

## Question 1

The variable `data` is a list of dictionaries.

Create a function `replace_unknown` that:
- Receives as argument a list of dictionaries;
- In each dictionary, it replaces all the '?' you find with the string value '1';
- Returns the modified list of dictionaires;

Note: values should retain the original datatypes.

<!-- grade_id: cell-7d5d16f632d819c3, points: 0, grade: false, solution: true -->
```python
def replace_unknown(data):
    lista=[]
    for dicio in data:
        dicio2={}
        for k, v in dicio.items():
            if v =='?':
            #k_= str(k).replace('?', '1')  
                v_= str(v).replace('?', '1')
                dicio2[k] = v_
            else:            
                dicio2[k] = v     
        lista.append(dicio2)
    return lista
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-2ec0be01019aa1ce, points: 3, grade: true, solution: false -->
```python
utils.b2w2_exerc_1_grading(replace_unknown)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 2

The variable `data` is a list of dictionaries.

Create a function `convert_int` that:
- Receives as argument a list of dictionaries;
- In each dictionary, find all the numeric values and make sure their data type is `int`;
- Returns the modified list of dictionaires;

Hint: `isdigit()`

<!-- grade_id: cell-5b768d0072b87818, points: 0, grade: false, solution: true -->
```python
def convert_int(data):
    lista=[]
    for dicio in data:
        dicio2={}
        for k,v in dicio.items():
            if type(v)==str and v.isdigit():
                v=int(v)
            dicio2[k]=v    
        lista.append(dicio2)
    return lista        
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-9138a4e7a89f3adc, points: 3, grade: true, solution: false -->
```python
utils.b2w2_exerc_2_grading(convert_int)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 3

The variable `data` is a list of dictionaries.

Create a function `lower_class` that:

- Receives as argument a list of dictionaries;
- For each element of the list, find the key 'class_type' and convert the correspondent value to lower case;
- Returns the modified list of dictionaires;

<!-- grade_id: cell-0f7774d47c55eb30, points: 0, grade: false, solution: true -->
```python
def lower_class(data):
    lista=[]
    for dicio in data:
        dicio2={}
        for k,v in dicio.items():            
            if k == 'class_type':
                v=v.lower()
            dicio2[k]=v    
        lista.append(dicio2)
    return lista 
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-3a68fa6f6e905977, points: 3, grade: true, solution: false -->
```python
utils.b2w2_exerc_3_grading(lower_class)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 4

The variable `data` is a list of dictionaries.

Create a function `order_animals` that:

- Receives as argument a list of dictionaries;
- Order the animals by number of legs;
- Return them as a list of (animal_name, class_type, legs) tuples;

<!-- grade_id: cell-122a9ec5cefa138a, points: 0, grade: false, solution: true -->
```python
def order_animals(data):
    lista=[]
    for dicio in data:
        lista2=[]
        for k,v in dicio.items(): 
            if k == 'animal_name':
                lista2.append(v)
            if k == 'class_type':
                lista2.append(v) 
            if k == 'legs':
                lista2.append(v)     
        lista.append(lista2)
    
    sorted_list= sorted(lista, key=lambda x: int(x[1]))
    swept=list(map(lambda x: (x[0], x[2], x[1]), sorted_list))
    
    return swept        
                
                
    # YOUR CODE HERE
    #raise NotImplementedError()
```

```python
ab=order_animals(data)
```

```python
ab
```

Output:
```
[('haddock', 'Fish', 0),
 ('gull', 'Bird', 2),
 ('gull', 'Bird', 2),
 ('pheasant', 'Bird', 2),
 ('crow', 'Bird', 2),
 ('wren', 'Bird', 2),
 ('hawk', 'Bird', 2),
 ('hare', 'Mammal', 4),
 ('cavy', 'Mammal', 4),
 ('calf', 'Mammal', 4),
 ('giraffe', 'Mammal', 4)]
```

<!-- grade_id: cell-4c4b419744a6ebf6, points: 3, grade: true, solution: false -->
```python
utils.b2w2_exerc_4_grading(order_animals)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 5

The variable `data` is a list of dictionaries.

Create a function `mammal_letters` that:

- Receives as argument a list of dictionaries;
- Filter all the mammals;
- Return the average number of letters in the mammals' names. The return type must be a float;

<!-- grade_id: cell-81f72032281eb943, points: 0, grade: false, solution: true -->
```python
def mammals_letters(data):
    lista=[]
    for dicio in data:
        for k, v in dicio.items():
            if v=='mammal':
                name=dicio['animal_name']
                lista.append(len(name))
                
    count=0
    for nr in lista:
        count += nr
        #print(count)
    return count/len(lista)
        
    # YOUR CODE HERE
    #raise NotImplementedError()
```

```python
abc= mammals_letters(data)
```

```python
abc
```

Output:
```
19
```

```python
lista=[]
for dicio in data:
    for k, v in dicio.items():
        if v=='mammal':
            name=dicio['animal_name']
            lista.append(len(name))
```

```python
lista
```

Output:
```
[4, 4, 4, 7]
```

```python
data_lower = utils.generate_data_lower()
```

```python
data_lower
```

Output:
```
[{'animal_name': 'gull',
  'hair': 0,
  'feathers': 1,
  'eggs': 1,
  'milk': 0,
  'aquatic': 1,
  'toothed': 0,
  'fins': 0,
  'legs': 2,
  'class_type': 'bird'},
 {'animal_name': 'hare',
  'hair': 1,
  'feathers': 0,
  'eggs': 0,
  'milk': 1,
  'aquatic': 0,
  'toothed': 1,
  'fins': 0,
  'legs': 4,
  'class_type': 'mammal'},
 {'animal_name': 'gull',
  'hair': 0,
  'feathers': 1,
  'eggs': 1,
  'milk': 0,
  'aquatic': 1,
  'toothed': 0,
  'fins': 0,
  'legs': 2,
  'class_type': 'bird'},
 {'animal_name': 'pheasant',
  'hair': 0,
  'feathers': 1,
  'eggs': 1,
  'milk': 0,
  'aquatic': 0,
  'toothed': 0,
  'fins': 0,
  'legs': 2,
  'class_type': 'bird'},
 {'animal_name': 'cavy',
  'hair': 1,
  'feathers': 0,
  'eggs': 0,
  'milk': 1,
  'aquatic': 0,
  'toothed': 1,
  'fins': 0,
  'legs': 4,
  'class_type': 'mammal'},
 {'animal_name': 'haddock',
  'hair': 0,
  'feathers': 0,
  'eggs': 1,
  'milk': 0,
  'aquatic': 1,
  'toothed': 1,
  'fins': 1,
  'legs': 0,
  'class_type': 'fish'},
 {'an
```

<!-- grade_id: cell-03842b58f0727f29, points: 3, grade: true, solution: false -->
```python
utils.b2w2_exerc_5_grading(mammals_letters)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 6

Create an `Animal` class.

The class has three class parameters: `animal_name` (str), `animal_type` (str) and `eggs` (bool).
Implement a method called `lay_eggs` that:
- Returns random number between 0 and 10 if the animal can lay eggs (eggs parameter = True)
- Returns the sentence "I can't lay eggs..." if the animal doesn't lay eggs

<!-- grade_id: cell-786ac51bf37eef4c, points: 0, grade: false, solution: true -->
```python
class Animal:
    def __init__(self, animal_name: str, animal_type: str, eggs: bool ):
        self.animal_name = animal_name
        self.animal_type= animal_type
        self.eggs=eggs
        
    def lay_eggs(self):
        if self.eggs:
            return random.randrange(0,9)
        else:
            return "I can't lay eggs..."
            

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-f50a68ec0e515188, points: 5, grade: true, solution: false -->
```python
utils.b2w2_exerc_6_grading(Animal)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

### Last but not least, submit your work!

To submit your work, fill your slack ID in the `slack_id` variable (as a string).

Example: `slack_id = "x-men"`

Help: if you forgot your slack ID, [read this](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c).

<!-- grade_id: cell-1a02f44222e2455e, points: 0, grade: false, solution: true -->
```python
# Submit your work!

slack_id = 'U01SZNW1QAD'

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-c461aefc9884225e, points: 0, grade: true, solution: false -->
```python
from submit import submit
assert isinstance(slack_id, str)

slu = 17_3
submit(slack_id, slu)
```

```python

```