# SLU 17 - Exam Prep I

## Batch 1 Python Exam 
This is the Python exam from batch 1. Exam Duration: 2h

```python
import json
#used for evaluation
import utils
```

## Question 1

Complete the function `import_data`. This function should:
- Open the data/countries.json file;
- Use Python’s module json to load the data from the file into a dictionary;
- Return the dictionary;

<!-- grade_id: cell-ce9179c0612fd6b0, points: 0, grade: false, solution: true -->
```python
def import_data():
    f= open('data/countries.json')  
    data=json.load(f)
    f.close
    return data
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-78d54c806fb28266, points: 2, grade: true, solution: false -->
```python
utils.b1_exerc_1_grading(import_data)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 2

The dictionary that you built contains information about a series of countries. So from now on, we’ll call it the `countries` dictionary.
In the `countries` dictionary, each country has a field called `Population` that gives the number of people living in that country. Since this is a count, the data should have integer type. However, we have float…

Complete the function `convert_population_to_int`. This function should:
- Receive the `countries` dictionary as argument;
- Convert the `Population` field from float to int type for all the countries, using a for loop;
- Return the updated `countries` dictionary;

```python
countries=import_data()
```

```python
for key,value in countries.items():
    ab=int(value['Population'])
    countries[key]['Population']=ab
countries    
```

Output:
```
{'Afghanistan': {'Continent': 'Asia',
  'Population': 31056997,
  'Area': 647500,
  'GDP': 700.0,
  'Literacy': 36.0,
  'Phones': 3.22},
 'Albania': {'Continent': 'Europe',
  'Population': 3581655,
  'Area': 28748,
  'GDP': 4500.0,
  'Literacy': 86.5,
  'Phones': 71.2},
 'Algeria': {'Continent': 'Africa',
  'Population': 32930091,
  'Area': 2381740,
  'GDP': 6000.0,
  'Literacy': 70.0,
  'Phones': 78.1},
 'American Samoa': {'Continent': 'Oceania',
  'Population': 57794,
  'Area': 199,
  'GDP': 8000.0,
  'Literacy': 97.0,
  'Phones': 259.54},
 'Andorra': {'Continent': 'Europe',
  'Population': 71201,
  'Area': 468,
  'GDP': 19000.0,
  'Literacy': 100.0,
  'Phones': 497.18},
 'Angola': {'Continent': 'Africa',
  'Population': 12127071,
  'Area': 1246700,
  'GDP': 1900.0,
  'Literacy': 42.0,
  'Phones': 7.78},
 'Anguilla': {'Continent': 'America',
  'Population': 13477,
  'Area': 102,
  'GDP': 8600.0,
  'Literacy': 95.0,
  'Phones': 460.04},
 'Antigua & Barbuda': {'Continent': 'America',
 
```

<!-- grade_id: cell-28d89deeef261406, points: 0, grade: false, solution: true -->
```python
def convert_population_to_int(countries):
    for key,value in countries.items():
        ab=int(value['Population'])
        countries[key]['Population']=ab
    return countries    
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-27243f1f1f55589c, points: 2, grade: true, solution: false -->
```python
utils.b1_exerc_2_grading(convert_population_to_int)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 3

In the countries dictionary, each country has a field called `Area`. This is the area of the country given in square miles.

Complete the function `convert_area_to_sq_km`. This function should:
- Receive the `countries` dictionary as argument;
- Convert the `Area` field from square miles to square kilometres;
 - Use the following conversion: 1 sq mi = 2.58999 sq km;
 - Round the result to 1 decimal digit;
- Return the updated `countries` dictionary;

<!-- grade_id: cell-3fd27539b41bed5c, points: 0, grade: false, solution: true -->
```python
def convert_area_to_sq_km(countries):
    for key, value in countries.items():
        area=value['Area']*2.58999
        countries[key]['Area']= round(area, 1)
    return countries    
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-b77708fde98d925a, points: 2, grade: true, solution: false -->
```python
utils.b1_exerc_3_grading(convert_area_to_sq_km)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 4

Complete the function `get_europe_countries`. This function should:
- Receive the `countries` dictionary as argument;
- Build a list with the names of the countries in Europe sorted by alphabetical order (from A to Z);
 - Use the field called `Continent`;
 - Use list comprehension;
 - Use the sort method from the list data type;
- Return the list;

<!-- grade_id: cell-ee1f38ef730c2004, points: 0, grade: false, solution: true -->
```python
def get_europe_countries(countries):
    lista=[k for k,v in countries.items() if v['Continent']== 'Europe']
    lista.sort()
    return lista
    # YOUR CODE HERE
    #raise NotImplementedError()
```

```python
abc1=[k for k,v in countries.items() if v['Continent']== 'Europe']
abc.sort()
print(abc1)
```

Output:
```
['Albania', 'Andorra', 'Austria', 'Belgium', 'Bulgaria', 'Croatia', 'Czech Republic', 'Denmark', 'Estonia', 'Finland', 'France', 'Germany', 'Greece', 'Hungary', 'Iceland', 'Ireland', 'Italy', 'Latvia', 'Liechtenstein', 'Lithuania', 'Luxembourg', 'Malta', 'Monaco', 'Netherlands', 'Norway', 'Poland', 'Portugal', 'Romania', 'San Marino', 'Serbia', 'Slovenia', 'Spain', 'Sweden', 'Switzerland', 'United Kingdom']
```

<!-- grade_id: cell-48bf94a05b056aa0, points: 2, grade: true, solution: false -->
```python
utils.b1_exerc_4_grading(get_europe_countries)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 5

Complete the function `get_literacy_levels_by_continent`. This function should:
- Receive the `countries` dictionary as argument;
- Receive a continent name as argument;
- For the countries in the continent (received as argument), compute the literacy level as:
 - literacy in [0, 25[ % - VERY_LOW;
 - literacy in [25, 50[ % - LOW;
 - literacy in [50, 70[ % - MEDIUM;
 - literacy in [70, 90[ % - HIGH;
 - literacy in [90, 100] % - VERY_HIGH;
- Return a list of tuples like: [(country_1, literacy_1, literacy_level_1), ..., (country_n, literacy_n, literacy_level_n), ...];

<!-- grade_id: cell-3911f29ed53cd968, points: 0, grade: false, solution: true -->
```python
def get_literacy_levels_by_continent(countries, continent):
    lista=[]
    for k,v in countries.items():
        if v['Continent'] == continent:
            if 0 <= v['Literacy'] < 25:
                lista.append((k, v['Literacy'], 'VERY_LOW'))
            elif 25 <= v['Literacy'] < 50:    
                lista.append((k, v['Literacy'], 'LOW'))
            elif 50 <= v['Literacy'] < 70:    
                lista.append((k, v['Literacy'], 'MEDIUM'))
            elif 70<= v['Literacy'] < 90:    
                lista.append((k, v['Literacy'], 'HIGH'))
            else:    
                lista.append((k, v['Literacy'], 'VERY_HIGH'))
    return lista
    
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-fa41a51f9ff00d7c, points: 2, grade: true, solution: false -->
```python
utils.b1_exerc_5_grading(get_literacy_levels_by_continent, "Africa")
utils.b1_exerc_5_grading(get_literacy_levels_by_continent, "Europe")
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 6

Complete the function `get_country_codes`. This function should:
- Receive the `countries` dictionary as argument;
- Build a list with all the countries names;
- Using `map` and a `lambda` function, convert the country names in country codes by selecting the 3 first letters of the country;
- Returns this list;

Note: Don't worry about duplicates on the result.

<!-- grade_id: cell-bd068242631b4909, points: 0, grade: false, solution: true -->
```python
def get_country_codes(countries):
    lista=[]
    x_lambda=lambda a: a[:3]
    for k in countries.keys():
        lista.append(k)
    #codes=[]    
    #for country in lista:      
    codes= list(map(x_lambda, lista))
    return codes
    # YOUR CODE HERE
    #raise NotImplementedError()
```

```python
lista=[]

for k in countries.keys():
        #x=lambda a: a[:3]
        lista.append(k)
        
lista        
```

Output:
```
['Afghanistan',
 'Albania',
 'Algeria',
 'American Samoa',
 'Andorra',
 'Angola',
 'Anguilla',
 'Antigua & Barbuda',
 'Argentina',
 'Aruba',
 'Australia',
 'Austria',
 'Bahamas, The',
 'Bahrain',
 'Bangladesh',
 'Barbados',
 'Belgium',
 'Belize',
 'Benin',
 'Bermuda',
 'Bhutan',
 'Bolivia',
 'Botswana',
 'Brazil',
 'British Virgin Is.',
 'Brunei',
 'Bulgaria',
 'Burkina Faso',
 'Burma',
 'Burundi',
 'Cambodia',
 'Cameroon',
 'Canada',
 'Cape Verde',
 'Cayman Islands',
 'Central African Rep.',
 'Chad',
 'Chile',
 'China',
 'Colombia',
 'Comoros',
 'Congo, Dem. Rep.',
 'Congo, Repub. of the',
 'Cook Islands',
 'Costa Rica',
 "Cote d'Ivoire",
 'Croatia',
 'Cuba',
 'Czech Republic',
 'Denmark',
 'Djibouti',
 'Dominica',
 'Dominican Republic',
 'Ecuador',
 'Egypt',
 'El Salvador',
 'Equatorial Guinea',
 'Eritrea',
 'Estonia',
 'Ethiopia',
 'Fiji',
 'Finland',
 'France',
 'French Guiana',
 'French Polynesia',
 'Gabon',
 'Gambia, The',
 'Germany',
 'Ghana',
 'Greece',
 'Grenada',
 'Guadeloupe
```

```python
x_lambda = lambda a: a[:3] 
abc=list(map(x_lambda, lista))
abc
```

Output:
```
['Afg',
 'Alb',
 'Alg',
 'Ame',
 'And',
 'Ang',
 'Ang',
 'Ant',
 'Arg',
 'Aru',
 'Aus',
 'Aus',
 'Bah',
 'Bah',
 'Ban',
 'Bar',
 'Bel',
 'Bel',
 'Ben',
 'Ber',
 'Bhu',
 'Bol',
 'Bot',
 'Bra',
 'Bri',
 'Bru',
 'Bul',
 'Bur',
 'Bur',
 'Bur',
 'Cam',
 'Cam',
 'Can',
 'Cap',
 'Cay',
 'Cen',
 'Cha',
 'Chi',
 'Chi',
 'Col',
 'Com',
 'Con',
 'Con',
 'Coo',
 'Cos',
 'Cot',
 'Cro',
 'Cub',
 'Cze',
 'Den',
 'Dji',
 'Dom',
 'Dom',
 'Ecu',
 'Egy',
 'El ',
 'Equ',
 'Eri',
 'Est',
 'Eth',
 'Fij',
 'Fin',
 'Fra',
 'Fre',
 'Fre',
 'Gab',
 'Gam',
 'Ger',
 'Gha',
 'Gre',
 'Gre',
 'Gua',
 'Gua',
 'Gua',
 'Gui',
 'Gui',
 'Guy',
 'Hai',
 'Hon',
 'Hon',
 'Hun',
 'Ice',
 'Ind',
 'Ind',
 'Ira',
 'Ira',
 'Ire',
 'Isr',
 'Ita',
 'Jam',
 'Jap',
 'Jor',
 'Ken',
 'Kor',
 'Kor',
 'Kuw',
 'Lao',
 'Lat',
 'Leb',
 'Les',
 'Lib',
 'Lib',
 'Lie',
 'Lit',
 'Lux',
 'Mac',
 'Mad',
 'Mal',
 'Mal',
 'Mal',
 'Mal',
 'Mal',
 'Mar',
 'Mar',
 'Mau',
 'Mau',
 'Mex',
 'Mic',
 'Mon',
 'Mon',
 'Mor',
 'Moz',
 'Nam',
 'Nep',
 'Net',

```

<!-- grade_id: cell-5f1eaf642616c8a1, points: 2, grade: true, solution: false -->
```python
utils.b1_exerc_6_grading(get_country_codes)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 7

The `Country` class is going to be used to represent countries and to return information about them.

### Question 7.1

Complete the `get_population_in_millions` method. This method should:
- Return the population of the country that the class represents;
- The population should be returned in millions with 2 decimal digits;

<!-- grade_id: cell-9b9ae4ba8e61e508, points: 0, grade: false, solution: true -->
```python
class Country:
    def __init__(self, country_name, population, continent):
        self.country_name = country_name
        self.population = population
        self.continent = continent

    def get_population_in_millions(self):
        self.population_mil=round(self.population/1e6, 2)
        return self.population_mil
        # YOUR CODE HERE
        #raise NotImplementedError()
```

<!-- grade_id: cell-9e28e403f317f23d, points: 2, grade: true, solution: false -->
```python
utils.b1_exerc_7_1_grading(Country("Penguinea", 12345678, "Antarctica"), Country)
utils.b1_exerc_7_1_grading(Country("Walrussia", 8344567, "Antarctica"), Country)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 7.2

Complete the `get_country_population`. This function should:
- Receive the `countries` dictionary as argument;
- Receive a country name as argument;
- Try to find the country (received as argument) in the countries dictionary;
- If the country doesn’t exist in the countries dictionary, catch the KeyError, print a statement with the information that there is no information for that country and return `None`;
- If the country exists, create an object of the `Country` class;
- Call the `get_population_in_millions` method for the country;
- Return the population in millions value;

<!-- grade_id: cell-803b263dbee2e300, points: 0, grade: false, solution: true -->
```python
def get_country_population(countries, country_name):
    try:
        assert countries[country_name], 'There is no information for that country!'
        pais=Country(country_name, countries[country_name]['Population'], countries[country_name]['Continent'])
        return pais.get_population_in_millions()
    except KeyError:
        return None
    
    
    
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-0d0db73f5256c8e5, points: 6, grade: true, solution: false -->
```python
utils.b1_exerc_7_2_grading("Turkey", get_country_population)
utils.b1_exerc_7_2_grading("Narnia", get_country_population)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

### Last but not least, submit your work!

To submit your work, fill your slack ID in the `slack_id` variable (as a string).

Example: `slack_id = "x-men"`

Help: if you forgot your slack ID, [read this](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c).

<!-- grade_id: cell-6d2bcecf8ecbc46d, points: 0, grade: false, solution: true -->
```python
# Submit your work!

slack_id = 'U01SZNW1QAD'

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-f66c1a17d9cccc05, points: 0, grade: true, solution: false -->
```python
from submit import submit
assert isinstance(slack_id, str)

slu = 17_1
submit(slack_id, slu)
```

Output:
```
Success
```

```python

```